<problem display_name="Einon koetulokset" markdown="null">
  <script>
def test_percent(expect, ans, eps=1):
  ok = True
  msg = 'Oikein!'

  if ok:
    try:
      n_ans = float(ans.replace("%", ""))
    except ValueError:
      ok = False
      msg = 'Luvun on oltava desimaaliluku'
  if ok:
    if abs(n_ans - float(expect)*100) &gt; eps:
      ok = False
      msg = f'Väärin, tämä on nyt {n_ans} %'

  return { 'ok': ok, 'msg': msg }


EXPECT=[
  32/36,
  30/45,
  66/75,
  54/60,
  26/55
]
def test_mult_table(e, ans):
  res = []
  all_ok = True
  for i, v in enumerate(ans):
    expect = EXPECT[i]
    test_res = test_percent(expect, v)
    all_ok = all_ok and test_res['ok']
    res.append(test_res)
  return {
    'overall_message': 'Hyvä!' if all_ok else 'Vielä vähän korjattavaa!',
    'input_list': res
  }
</script>
  <p>Alla olevassa taulukossa ovat Einon koetulokset. Muunna koetulokset prosenttiluvuiksi.</p>
  <customresponse cfn="test_mult_table" expect="a">
    <style>
      #taulu_t1 th {
          text-align: center;
      }
      #taulu_t1 td {
          text-align: center;
      }
      #taulu_t1 tr {
          border-bottom: 1px solid black;
      }
      #taulu_t1 .tdh {
      	background-color: #CCCCCC;
        text-align: center;
      border-right: 1px solid black;
      }
      #taulu_t1 .tdin input {
          width: 50px;
      }
    </style>
    <table id="taulu_t1">
      <thead>
        <tr>
          <th>Aine</th>
          <th>Pisteet</th>
          <th>Maksimipistemäärä</th>
          <th>Prosenttiosuus</th>
        </tr>
      </thead>
      <tbody>
        <tr>
          <td class="tdh">Matematiikka</td>
          <td>\( 32 \)</td>
          <td>\( 36 \)</td>
          <td>
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Historia</td>
          <td class="tdin">\( 30 \)</td>
          <td class="tdin">\( 45 \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Biologia</td>
          <td class="tdin">\( 66 \)</td>
          <td class="tdin">\( 75 \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Englanti</td>
          <td class="tdin">\( 54 \)</td>
          <td class="tdin">\( 60 \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Uskonto</td>
          <td class="tdin">\( 26 \)</td>
          <td class="tdin">\( 55 \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
      </tbody>
    </table>
  </customresponse>
  <multiplechoiceresponse>
    <label>Missä aineessa hän pärjäsi parhaiten? </label>
    <choicegroup type="MultipleChoice">
      <choice correct="false">Matematiikka</choice>
      <choice correct="false">Historia</choice>
      <choice correct="false">Biologia</choice>
      <choice correct="true">Englanti</choice>
      <choice correct="false">Uskonto</choice>
    </choicegroup>
  </multiplechoiceresponse>
  <demandhint>
    <hint>Kyseessä on edelleen sama prosenttiosuustehtävä. Mikä on perusarvo ja mikä on osa?</hint>
    <hint>Miten saadut prosenttiluvut voi tulkita? Mitä ne kertovat Einon menestyksestä?</hint>
  </demandhint>
</problem>
