<problem display_name="Aurinkokunta" markdown="null">
  <script>
E_MASS=5.974e1
def test_percent(expect, ans, eps=0.1):
  ok = True
  msg = 'Oikein!'
  
  if ok:
    try:
      n_ans = float(ans.replace("%", ""))
    except ValueError:
      ok = False
      msg = 'Luvun on oltava desimaaliluku'
  if ok:
    if abs(n_ans - float(expect)*100) &gt; eps:
      ok = False
      res = f"{(float(n_ans)/100 * E_MASS * 1e23):.3e}"
      if "e" in res:
        res = res.replace("e", " \cdot 10^{").replace("+", "") + "}"
      msg = f'Väärin, planeetan massa olisi \( {res}\ \\textrm{{kg}} \)'

  return { 'ok': ok, 'msg': msg }

MASSES=[
  3.304,
  4.8745e1,
  5.974e1,
  6.422,
  1.899e4,
  5.684e3,
  8.686e2,
  1.029e3
]
EXPECT=[ m / E_MASS for m in MASSES ]
def test_mult_table(e, ans):
  res = []
  all_ok = True
  for i, v in enumerate(ans):
    expect = EXPECT[i]
    test_res = test_percent(expect, v)
    all_ok = all_ok and test_res['ok']
    res.append(test_res)
  return {
    'overall_message': 'Hyvä!' if all_ok else 'Vielä vähän korjattavaa!',
    'input_list': res
  }
</script>
  <p>Taulukossa on aurinkokuntamme planeettojen massat. Laske montako prosenttia kunkin planeetan
massa on Maan massasta. Ilmoita tulos prosentin kymmenesosan tarkkuudeella.</p>
  <customresponse cfn="test_mult_table" expect="a">
    <style>
      #taulu_t1 th {
          text-align: center;
      }
      #taulu_t1 td {
          text-align: center;
      }
      #taulu_t1 tr {
          border-bottom: 1px solid black;
      }
      #taulu_t1 .tdh {
      	background-color: #CCCCCC;
        text-align: center;
      border-right: 1px solid black;
      }
      #taulu_t1 .tdin input {
          width: 50px;
      }
    </style>
    <table id="taulu_t1">
      <thead>
        <tr>
          <th>Planeetta</th>
          <th>Massa \( [ \textrm{kg} ] \)</th>
          <th>Prosenttiosuus</th>
        </tr>
      </thead>
      <tbody>
        <tr>
          <td class="tdh">Merkurius</td>
          <td>\( 3.304 \cdot 10^{23} \)</td>
          <td>
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Venus</td>
          <td class="tdin">\( 4.8745 \cdot 10^{24} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Maa</td>
          <td class="tdin">\( 5.974 \cdot 10^{24} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Mars</td>
          <td class="tdin">\( 6.422 \cdot 10^{23} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Jupiter</td>
          <td class="tdin">\( 1.899 \cdot 10^{27} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Saturnus</td>
          <td class="tdin">\( 5.684 \cdot 10^{26} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Uranus</td>
          <td class="tdin">\( 8.686 \cdot 10^{25} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
        <tr>
          <td class="tdh">Neptunus</td>
          <td class="tdin">\( 1.029 \cdot 10^{26} \)</td>
          <td class="tdin">
            <textline size="10"/>
          </td>
        </tr>
      </tbody>
    </table>
  </customresponse>
</problem>
